// primitive_types1.rs
// 填写代码行的缺失部分
// 没有提示，没有技巧，只要习惯键入这些 :)

// I AM NOT DONE

fn main() {
    // Booleans (`bool`)

    let is_morning = true;
    if is_morning {
        println!("Good morning!");
    }

    let // 像示例一样完成这行！或者将其设为 false ！
    if is_evening {
        println!("Good evening!");
    }
}
