// primitive_types3.rs
// Create an array with at least 100 elements in it where the ??? is.
// 执行 `rustlings hint primitive_types3` 或在观察模式下使用 `hint` 子命令来获取提示。

// I AM NOT DONE

fn main() {
    let a = ???

    if a.len() >= 100 {
        println!("Wow, that's a big array!");
    } else {
        println!("Meh, I eat arrays like that for breakfast.");
    }
}
