// primitive_types5.rs
// Destructure the `cat` tuple so that the println will work.
// 执行 `rustlings hint primitive_types5` 或在观察模式下使用 `hint` 子命令来获取提示。

// I AM NOT DONE

fn main() {
    let cat = ("Furry McFurson", 3.5);
    let /* your pattern here */ = cat;

    println!("{} is {} years old.", name, age);
}
